<!DOCTYPE html>
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<html lang="en-US" dir="ltr">
<!--******************************************-->
<!--*  Authored with PreTeXt                 *-->
<!--*  pretextbook.org                       *-->
<!--*  Theme: default-modern                 *-->
<!--*  Palette: blue-red                     *-->
<!--******************************************-->
<head xmlns:og="http://ogp.me/ns#" xmlns:book="https://ogp.me/ns/book#">
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<title>SB ActiveCode</title>
<link rel="preconnect" href="https://fonts.googleapis.com">
<link rel="preconnect" href="https://fonts.gstatic.com" crossorigin="">
<link rel="stylesheet" href="https://fonts.googleapis.com/css2?family=Material+Symbols+Outlined:opsz,wght,FILL,GRAD@24,400,0,0">
<link href="_static/pretext/css/theme.css" rel="stylesheet" type="text/css">
<link href="_static/pretext/css/ol-markers.css" rel="stylesheet" type="text/css">
<script>
var runestoneMathReady = new Promise((resolve) => window.rsMathReady = resolve);
window.MathJax = {
  "tex": {
    "inlineMath": [
      [
        "\\(",
        "\\)"
      ]
    ],
    "tags": "none",
    "tagSide": "right",
    "tagIndent": ".8em",
    "packages": {
      "[+]": [
        "base",
        "extpfeil",
        "ams",
        "amscd",
        "color",
        "newcommand",
        "knowl"
      ]
    }
  },
  "options": {
    "ignoreHtmlClass": "tex2jax_ignore|ignore-math",
    "processHtmlClass": "process-math",
    "renderActions": {
      "findScript": [
        10,
        function (doc) {
            document.querySelectorAll('script[type^="math/tex"]').forEach(function(node) {
                var display = !!node.type.match(/; *mode=display/);
                var math = new doc.options.MathItem(node.textContent, doc.inputJax[0], display);
                var text = document.createTextNode('');
                node.parentNode.replaceChild(text, node);
                math.start = {node: text, delim: '', n: 0};
                math.end = {node: text, delim: '', n: 0};
                doc.math.push(math);
            });
        },
        ""
      ]
    }
  },
  "chtml": {
    "scale": 0.98,
    "mtextInheritFont": true
  },
  "loader": {
    "load": [
      "input/asciimath",
      "[tex]/extpfeil",
      "[tex]/amscd",
      "[tex]/color",
      "[tex]/newcommand",
      "[pretext]/mathjaxknowl3.js"
    ],
    "paths": {
      "pretext": "_static/pretext/js/lib"
    }
  },
  "startup": {
    pageReady() {
      return MathJax.startup.defaultPageReady().then(function () {
      console.log("in ready function");
      rsMathReady();
      }
    )}
  }
};
</script><script src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-chtml.js"></script><meta name="Keywords" content="abstract algebra, PreTeXt, samples, Authored in PreTeXt">
<meta name="viewport" content="width=device-width, initial-scale=1.0">
<meta property="og:type" content="book">
<meta property="og:image" content="https://pretextbook.org/examples/noparts/decorative/html/external/cover_aata_2014.png">
<meta property="book:title" content="PreTeXt Sample Book">
<meta property="book:author" content="Thomas W. Judson">
<script src="_static/pretext/js/lib/jquery.min.js"></script><script src="_static/pretext/js/lib/jquery.sticky.js"></script><script src="_static/pretext/js/lib/jquery.espy.min.js"></script><script src="_static/pretext/js/pretext.js"></script><script src="_static/pretext/js/pretext_add_on.js?x=1"></script><script src="_static/pretext/js/user_preferences.js"></script><!--** eBookCongig is necessary to configure interactive       **-->
<!--** Runestone components to run locally in reader's browser **-->
<!--** No external communication:                              **-->
<!--**     log level is 0, Runestone Services are disabled     **-->
<script type="text/javascript">
eBookConfig = {};
eBookConfig.useRunestoneServices = false;
eBookConfig.host = 'http://127.0.0.1:8000';
eBookConfig.course = 'PTX_Course_Title_Here';
eBookConfig.basecourse = 'PTX_Base_Course';
eBookConfig.isLoggedIn = false;
eBookConfig.email = '';
eBookConfig.isInstructor = false;
eBookConfig.logLevel = 0;
eBookConfig.username = '';
eBookConfig.readings = null;
eBookConfig.activities = null;
eBookConfig.downloadsEnabled = false;
eBookConfig.allow_pairs = false;
eBookConfig.enableScratchAC = true;
eBookConfig.acDefaultLanguage = 'python';
eBookConfig.build_info = "";
eBookConfig.python3 = null;
eBookConfig.runestone_version = '7.8.4';
eBookConfig.jobehost = '';
eBookConfig.proxyuri_runs = '';
eBookConfig.proxyuri_files = '';
eBookConfig.enable_chatcodes =  false;
</script>
<!--*** Runestone Services ***-->
<script src="_static/prefix-runtime.962d00926ef9460e.bundle.js"></script><script src="_static/prefix-723.3e6434f80549315a.bundle.js"></script><script src="_static/prefix-runestone.2e7f40e86ae7d163.bundle.js"></script><link rel="stylesheet" type="text/css" href="_static/prefix-723.3bccd435914aa0ff.css">
<link rel="stylesheet" type="text/css" href="_static/prefix-runestone.5bdeacdba067a0e3.css">
<script src="https://sagecell.sagemath.org/static/embedded_sagecell.js"></script><script src="_static/pretext/js/lti_iframe_resizer.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/components/prism-core.min.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/autoloader/prism-autoloader.min.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/line-numbers/prism-line-numbers.min.js" integrity="sha512-dubtf8xMHSQlExGRQ5R7toxHLgSDZ0K7AunqPWHXmJQ8XyVIG19S1T95gBxlAeGOK02P4Da2RTnQz0Za0H0ebQ==" crossorigin="anonymous" referrerpolicy="no-referrer"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/line-highlight/prism-line-highlight.min.js" integrity="sha512-93uCmm0q+qO5Lb1huDqr7tywS8A2TFA+1/WHvyiWaK6/pvsFl6USnILagntBx8JnVbQH5s3n0vQZY6xNthNfKA==" crossorigin="anonymous" referrerpolicy="no-referrer"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/lunr.js/2.3.9/lunr.min.js" integrity="sha512-4xUl/d6D6THrAnXAwGajXkoWaeMNwEKK4iNfq5DotEbLPAfk6FSxSP3ydNxqDgCw1c/0Z1Jg6L8h2j+++9BZmg==" crossorigin="anonymous" referrerpolicy="no-referrer"></script><script src="lunr-pretext-search-index.js" async=""></script><script src="_static/pretext/js/pretext_search.js"></script><script src="_static/pretext/js/lib/knowl.js"></script><!--knowl.js code controls Sage Cells within knowls--><script>sagecellEvalName='Evaluate (Sage)';
</script>
</head>
<body id="PTXSB" class="pretext book ignore-math">
<a class="assistive" href="#ptx-content">Skip to main content</a><header id="ptx-masthead" class="ptx-masthead"><div class="ptx-banner">
<a id="logo-link" class="logo-link" target="_blank" href="http://abstract.pugetsound.edu"><img src="external/cover_aata_2014.png" alt="Logo image"></a><div class="title-container">
<h1 class="heading"><a href="sample-book.html"><span class="title">PreTeXt Sample Book</span> <span class="subtitle">Abstract Algebra (SAMPLE ONLY)</span></a></h1>
<p class="byline">Thomas W. Judson, Isaac Newton (Editor)</p>
</div>
</div></header><nav id="ptx-navbar" class="ptx-navbar navbar"><div class="ptx-navbar-contents">
<button class="toc-toggle button" title="Contents"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5d2;</span><span class="name">Contents</span></button><a class="index-button button" href="backmatter-8.html" title="Index"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe88e;</span><span class="name">Index</span></a><div class="searchbox">
<div class="searchwidget"><button id="searchbutton" class="searchbutton button" type="button" title="Search book"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe8b6;</span><span class="name">Search Book</span></button></div>
<div id="searchresultsplaceholder" class="searchresultsplaceholder" style="display: none">
<div class="search-results-controls">
<input aria-label="Search term" id="ptxsearch" class="ptxsearch" type="text" name="terms" placeholder="Search term"><button title="Close search" id="closesearchresults" class="closesearchresults"><span class="material-symbols-outlined">close</span></button>
</div>
<h2 class="search-results-heading">Search Results: </h2>
<div id="searchempty" class="searchempty"><span>No results.</span></div>
<ol id="searchresults" class="searchresults"></ol>
</div>
</div>
<span class="nav-other-controls"><button id="light-dark-button" class="light-dark-button button" title="Dark Mode"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe51c;</span><span class="name">Dark Mode</span></button></span><span class="treebuttons"><a class="previous-button button" href="inactive-code.html" title="Previous"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5cb;</span><span class="name">Prev</span></a><a class="up-button button" href="rune.html" title="Up"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5ce;</span><span class="name">Up</span></a><a class="next-button button" href="code-lens.html" title="Next"><span class="name">Next</span><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5cc;</span></a></span><span class="nav-runestone-controls"><button onclick="runestoneComponents.popupScratchAC()" class="activecode-toggle button" title="Open Scratch ActiveCode"><span class="icon material-symbols-outlined" aria-hidden="true">&#xf097;</span><span class="name">Scratch ActiveCode</span></button></span>
</div></nav><div id="latex-macros" class="hidden-content process-math" style="display:none"><span class="process-math">\(\newcommand{\identity}{\mathrm{id}}
\newcommand{\notdivide}{{\not{\mid}}}
\newcommand{\notsubset}{\not\subset}
\newcommand{\lcm}{\operatorname{lcm}}
\newcommand{\gf}{\operatorname{GF}}
\newcommand{\inn}{\operatorname{Inn}}
\newcommand{\aut}{\operatorname{Aut}}
\newcommand{\Hom}{\operatorname{Hom}}
\newcommand{\cis}{\operatorname{cis}}
\newcommand{\chr}{\operatorname{char}}
\newcommand{\Null}{\operatorname{Null}}
\newcommand{\lt}{&lt;}
\newcommand{\gt}{&gt;}
\newcommand{\amp}{&amp;}
\definecolor{fillinmathshade}{gray}{0.9}
\newcommand{\fillinmath}[1]{\mathchoice{\colorbox{fillinmathshade}{$\displaystyle     \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\textstyle        \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptstyle      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptscriptstyle\phantom{\,#1\,}$}}}
\newcommand{\sfrac}[2]{{#1}/{#2}}
\)</span></div>
<div class="ptx-page">
<div id="ptx-sidebar" class="ptx-sidebar"><nav id="ptx-toc" class="ptx-toc depth3 parts focused" data-preexpanded-levels="0" data-max-levels="3"><ul class="structural toc-item-list contains-active">
<li class="toc-item toc-frontmatter">
<div class="toc-title-box"><a href="frontmatter.html" class="internal"><span class="title">Front Matter</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-colophon"><div class="toc-title-box"><a href="front-colophon.html" class="internal"><span class="title">Colophon</span></a></div></li>
<li class="toc-item toc-biography"><div class="toc-title-box"><a href="author-bio-TWJ.html" class="internal"><span class="title">Author Biography</span></a></div></li>
<li class="toc-item toc-biography"><div class="toc-title-box"><a href="author-bio-RAB.html" class="internal"><span class="title">About Robert A. Beezer</span></a></div></li>
<li class="toc-item toc-dedication"><div class="toc-title-box"><a href="dedication.html" class="internal"><span class="title">Dedication</span></a></div></li>
<li class="toc-item toc-acknowledgement"><div class="toc-title-box"><a href="acknowledgement.html" class="internal"><span class="title">Acknowledgements</span></a></div></li>
<li class="toc-item toc-preface"><div class="toc-title-box"><a href="preface.html" class="internal"><span class="title">Preface</span></a></div></li>
<li class="toc-item toc-preface">
<div class="toc-title-box"><a href="contributors.html" class="internal"><span class="title">Contributors to the 4<span class="process-math">\(^\mathrm{th}\)</span> Edition</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-contributors"><div class="toc-title-box"><a href="contributors.html#contributors-3" class="internal"><span class="title"></span></a></div></li></ul>
</li>
</ul>
</li>
<li class="toc-item toc-part">
<div class="toc-title-box"><a href="part-basics.html" class="internal"><span class="codenumber">I</span> <span class="title">Basics</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-chapter">
<div class="toc-title-box"><a href="sets.html" class="internal"><span class="codenumber">1</span> <span class="title">Preliminaries</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-note-on-proofs.html" class="internal"><span class="codenumber">1.1</span> <span class="title">A Short Note on Proofs</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-note-on-proofs.html#section-note-on-proofs-3" class="internal"><span class="codenumber">1.1.1</span> <span class="title">Some Cautions and Suggestions</span></a></div></li></ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-sets-and-equivalence-relations.html" class="internal"><span class="codenumber">1.2</span> <span class="title">Sets and Equivalence Relations</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-sets-and-equivalence-relations.html#section-sets-and-equivalence-relations-2" class="internal"><span class="codenumber">1.2.1</span> <span class="title">Set Theory</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-sets-and-equivalence-relations.html#section-sets-and-equivalence-relations-3" class="internal"><span class="codenumber">1.2.2</span> <span class="title">Cartesian Products and Mappings</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-sets-and-equivalence-relations.html#section-sets-and-equivalence-relations-4" class="internal"><span class="codenumber">1.2.3</span> <span class="title">Equivalence Relations and Partitions</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="sets-sage.html" class="internal"><span class="codenumber">1.3</span> <span class="title">Sage</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-3" class="internal"><span class="codenumber">1.3.1</span> <span class="title">Executing Sage Commands</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-4" class="internal"><span class="codenumber">1.3.2</span> <span class="title">Immediate Help</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-5" class="internal"><span class="codenumber">1.3.3</span> <span class="title">Annotating Your Work</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-6" class="internal"><span class="codenumber">1.3.4</span> <span class="title">Lists</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-7" class="internal"><span class="codenumber">1.3.5</span> <span class="title">Lists of Integers</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="sets-sage.html#sets-sage-8" class="internal"><span class="codenumber">1.3.6</span> <span class="title">Saving and Sharing Your Work</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-sets.html" class="internal"><span class="codenumber">1.4</span> <span class="title">Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="sets-sage-exercises.html" class="internal"><span class="codenumber">1.5</span> <span class="title">Sage Exercises</span></a></div></li>
<li class="toc-item toc-references"><div class="toc-title-box"><a href="sets-8.html" class="internal"><span class="codenumber">1.6</span> <span class="title">References and Suggested Readings</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-chapter">
<div class="toc-title-box"><a href="integers.html" class="internal"><span class="codenumber">2</span> <span class="title">The Integers</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-section"><div class="toc-title-box"><a href="section-math-induction.html" class="internal"><span class="codenumber">2.1</span> <span class="title">Mathematical Induction and Math in a Title <span class="process-math">\(A\notsubset B\)</span></span></a></div></li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-division-algorithm.html" class="internal"><span class="codenumber">2.2</span> <span class="title">The Division Algorithm</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-division-algorithm.html#section-division-algorithm-3" class="internal"><span class="codenumber">2.2.1</span> <span class="title">The Euclidean Algorithm</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-division-algorithm.html#section-division-algorithm-4" class="internal"><span class="codenumber">2.2.2</span> <span class="title">Prime Numbers</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-division-algorithm.html#section-division-algorithm-5" class="internal"><span class="codenumber">2.2.3</span> <span class="title">Historical Note</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="integers-sage.html" class="internal"><span class="codenumber">2.3</span> <span class="title">Sage</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="integers-sage.html#integers-sage-3" class="internal"><span class="codenumber">2.3.1</span> <span class="title">Division Algorithm</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="integers-sage.html#integers-sage-4" class="internal"><span class="codenumber">2.3.2</span> <span class="title">Greatest Common Divisor</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="integers-sage.html#integers-sage-5" class="internal"><span class="codenumber">2.3.3</span> <span class="title">Primes and Factoring</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-integers.html" class="internal"><span class="codenumber">2.4</span> <span class="title">Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="integers-7.html" class="internal"><span class="codenumber">2.5</span> <span class="title">Programming Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="integers-sage-exercises.html" class="internal"><span class="codenumber">2.6</span> <span class="title">Sage Exercises</span></a></div></li>
<li class="toc-item toc-references"><div class="toc-title-box"><a href="integers-9.html" class="internal"><span class="codenumber">2.7</span> <span class="title">References and Suggested Readings</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li class="toc-item toc-part contains-active">
<div class="toc-title-box"><a href="part-algebra.html" class="internal"><span class="codenumber">II</span> <span class="title">Algebra (and Runestone)</span></a></div>
<ul class="structural toc-item-list contains-active">
<li class="toc-item toc-chapter">
<div class="toc-title-box"><a href="groups.html" class="internal"><span class="codenumber">3</span> <span class="title">Groups</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-mod-n-sym.html" class="internal"><span class="codenumber">3.1</span> <span class="title">Integer Equivalence Classes and Symmetries</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-mod-n-sym.html#section-mod-n-sym-5" class="internal"><span class="codenumber">3.1.1</span> <span class="title">The Integers mod <span class="process-math">\(n\)</span></span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-mod-n-sym.html#section-mod-n-sym-6" class="internal"><span class="codenumber">3.1.2</span> <span class="title">Symmetries</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-groups-define.html" class="internal"><span class="codenumber">3.2</span> <span class="title">Definitions and Examples</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-groups-define.html#section-groups-define-2" class="internal"><span class="codenumber">3.2.1</span> <span class="title">Definition of a Group</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-groups-define.html#section-groups-define-3" class="internal"><span class="codenumber">3.2.2</span> <span class="title">Examples of Groups</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-groups-define.html#section-groups-define-4" class="internal"><span class="codenumber">3.2.3</span> <span class="title">Basic Properties of Groups</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-groups-define.html#section-groups-define-5" class="internal"><span class="codenumber">3.2.4</span> <span class="title">Historical Note</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-subgroups.html" class="internal"><span class="codenumber">3.3</span> <span class="title">Subgroups</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-subgroups.html#section-subgroups-2" class="internal"><span class="codenumber">3.3.1</span> <span class="title">Definitions and Examples</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="section-subgroups.html#section-subgroups-3" class="internal"><span class="codenumber">3.3.2</span> <span class="title">Some Subgroup Theorems</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="groups-sage.html" class="internal"><span class="codenumber">3.4</span> <span class="title">Sage</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="groups-sage.html#groups-sage-3" class="internal"><span class="codenumber">3.4.1</span> <span class="title">Integers mod n</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="groups-sage.html#groups-sage-4" class="internal"><span class="codenumber">3.4.2</span> <span class="title">Groups of symmetries</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="groups-sage.html#groups-sage-5" class="internal"><span class="codenumber">3.4.3</span> <span class="title">Quaternions</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="groups-sage.html#groups-sage-6" class="internal"><span class="codenumber">3.4.4</span> <span class="title">Subgroups</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-groups.html" class="internal"><span class="codenumber">3.5</span> <span class="title">Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="groups-9.html" class="internal"><span class="codenumber">3.6</span> <span class="title">Additional Exercises: Detecting Errors</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="groups-sage-exercises.html" class="internal"><span class="codenumber">3.7</span> <span class="title">Sage Exercises</span></a></div></li>
<li class="toc-item toc-references"><div class="toc-title-box"><a href="groups-11.html" class="internal"><span class="codenumber">3.8</span> <span class="title">References and Suggested Readings</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-chapter">
<div class="toc-title-box"><a href="cyclic.html" class="internal"><span class="codenumber">4</span> <span class="title">Cyclicity</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-section"><div class="toc-title-box"><a href="cyclic-4.html" class="internal"><span class="codenumber">4.1</span> <span class="title">Cyclic groups</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="cyclic-5.html" class="internal"><span class="codenumber">4.2</span> <span class="title">Subgroups of a Cyclic Group</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="cyclic-6.html" class="internal"><span class="codenumber">4.3</span> <span class="title">Cyclic Groups of Complex Numbers</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="cyclic-7.html" class="internal"><span class="codenumber">4.4</span> <span class="title">Large Powers of Integers</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-cyclic.html" class="internal"><span class="codenumber">4.5</span> <span class="title">Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="cyclic-9.html" class="internal"><span class="codenumber">4.6</span> <span class="title">Programming Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="cyclic-sage-exercises.html" class="internal"><span class="codenumber">4.7</span> <span class="title">Sage Exercises</span></a></div></li>
<li class="toc-item toc-references"><div class="toc-title-box"><a href="cyclic-11.html" class="internal"><span class="codenumber">4.8</span> <span class="title">References and Suggested Readings</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-chapter contains-active">
<div class="toc-title-box"><a href="rune.html" class="internal"><span class="codenumber">5</span> <span class="title">Runestone Testing</span></a></div>
<ul class="structural toc-item-list contains-active">
<li class="toc-item toc-section"><div class="toc-title-box"><a href="inactive-code.html" class="internal"><span class="codenumber">5.1</span> <span class="title">Programs</span></a></div></li>
<li class="toc-item toc-section active"><div class="toc-title-box"><a href="activecode.html" class="internal"><span class="codenumber">5.2</span> <span class="title">ActiveCode</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="code-lens.html" class="internal"><span class="codenumber">5.3</span> <span class="title">Code Lens</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="coding-exercises.html" class="internal"><span class="codenumber">5.4</span> <span class="title">Coding Exercises</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="data-files.html" class="internal"><span class="codenumber">5.5</span> <span class="title">Data Files</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="videos-in-runestone.html" class="internal"><span class="codenumber">5.6</span> <span class="title">YouTube Videos</span></a></div></li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="rune-nested-divisions.html" class="internal"><span class="codenumber">5.7</span> <span class="title">Deeper</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="rune-nested-divisions.html#rune-nested-divisions-3" class="internal"><span class="codenumber">5.7.1</span> <span class="title">Subsection One</span></a></div></li>
<li class="toc-item toc-subsection"><div class="toc-title-box"><a href="rune-nested-divisions.html#rune-nested-divisions-4" class="internal"><span class="codenumber">5.7.2</span> <span class="title">Subsection Two</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="true-false-exercises.html" class="internal"><span class="codenumber">5.8</span> <span class="title">True/False Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-groupwork.html" class="internal"><span class="codenumber">5.9</span> <span class="title">Multiple Choice Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="parsons-exercises.html" class="internal"><span class="codenumber">5.10</span> <span class="title">Parsons Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="horizontal-parsons-exercises.html" class="internal"><span class="codenumber">5.11</span> <span class="title">Horizontal Parsons Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="matching-exercises.html" class="internal"><span class="codenumber">5.12</span> <span class="title">Matching Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="clickable-area-exercises.html" class="internal"><span class="codenumber">5.13</span> <span class="title">Clickable Area Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="select-exercises.html" class="internal"><span class="codenumber">5.14</span> <span class="title">Select Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="short-answer-exercises.html" class="internal"><span class="codenumber">5.15</span> <span class="title">Short Answer Exercises</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="polling.html" class="internal"><span class="codenumber">5.16</span> <span class="title">Polling</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="doenet.html" class="internal"><span class="codenumber">5.17</span> <span class="title">DoenetML</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="fill-in-exercises-old.html" class="internal"><span class="codenumber">5.18</span> <span class="title">Fill-In Exercises (Old Style)</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="fill-in-exercises.html" class="internal"><span class="codenumber">5.19</span> <span class="title">Fill-In Exercises</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="runestone-hodgepodge.html" class="internal"><span class="codenumber">5.20</span> <span class="title">Hodgepodge</span></a></div></li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="exercises-timed.html" class="internal"><span class="codenumber">5.21</span> <span class="title">Exercises that are Timed</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-timed.html#timed-exercises" class="internal"><span class="codenumber">5.21</span> <span class="title">Timed Exercises</span></a></div></li></ul>
</li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="interactive-projects.html" class="internal"><span class="codenumber">5.22</span> <span class="title">Projects and Friends</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="rune-25.html" class="internal"><span class="codenumber">5.23</span> <span class="title">Expedited Samples</span></a></div></li>
<li class="toc-item toc-reading-questions"><div class="toc-title-box"><a href="rune-26.html" class="internal"><span class="codenumber">5.24</span> <span class="title">Reading Questions</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="atomic-video.html" class="internal"><span class="codenumber">5.25</span> <span class="title">YouTube Video Embedding</span></a></div></li>
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="section-runestone-assignment-testing.html" class="internal"><span class="codenumber">5.26</span> <span class="title">Runestone Assignment Testing</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-exercises"><div class="toc-title-box"><a href="section-runestone-assignment-testing.html#section-runestone-assignment-testing-3" class="internal"><span class="codenumber">5.26</span> <span class="title">Exercises</span></a></div></li></ul>
</li>
<li class="toc-item toc-worksheet"><div class="toc-title-box"><a href="worksheet-groupwork.html" class="internal"><span class="codenumber">5.27</span> <span class="title">A “Group Work” Worksheet</span></a></div></li>
<li class="toc-item toc-exercises"><div class="toc-title-box"><a href="exercises-title-groupwork.html" class="internal"><span class="codenumber">5.28</span> <span class="title">Group Exercises</span></a></div></li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="splice-integration.html" class="internal"><span class="codenumber">5.29</span> <span class="title">SPLICE Integration</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li class="toc-item toc-backmatter">
<div class="toc-title-box"><a href="backmatter.html" class="internal"><span class="title">Back Matter</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-appendix"><div class="toc-title-box"><a href="backmatter-2.html" class="internal"><span class="codenumber">A</span> <span class="title">Notation</span></a></div></li>
<li class="toc-item toc-solutions"><div class="toc-title-box"><a href="backmatter-3.html" class="internal"><span class="codenumber">B</span> <span class="title">Hints and Answers to Selected Odd Exercises</span></a></div></li>
<li class="toc-item toc-solutions"><div class="toc-title-box"><a href="backmatter-4.html" class="internal"><span class="codenumber">C</span> <span class="title">Hints and Answers to Selected Even Exercises</span></a></div></li>
<li class="toc-item toc-solutions"><div class="toc-title-box"><a href="backmatter-5.html" class="internal"><span class="codenumber">D</span> <span class="title">Hints and Answers to Selected Exercises</span></a></div></li>
<li class="toc-item toc-appendix">
<div class="toc-title-box"><a href="backmatter-6.html" class="internal"><span class="codenumber">E</span> <span class="title">A Structured Appendix</span></a></div>
<ul class="structural toc-item-list">
<li class="toc-item toc-section">
<div class="toc-title-box"><a href="backmatter-6-3.html" class="internal"><span class="codenumber">E.1</span> <span class="title">A Section in an Appendix</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-subsection">
<div class="toc-title-box"><a href="backmatter-6-3.html#backmatter-6-3-3" class="internal"><span class="codenumber">E.1.1</span> <span class="title">A Subsection in a Section in an Appendix</span></a></div>
<ul class="structural toc-item-list"><li class="toc-item toc-subsubsection"><div class="toc-title-box"><a href="backmatter-6-3.html#backmatter-6-3-3-3" class="internal"><span class="codenumber">E.1.1.1</span> <span class="title">A Subsubsection in a Subsection in a Section in an Appendix</span></a></div></li></ul>
</li></ul>
</li>
<li class="toc-item toc-section"><div class="toc-title-box"><a href="backmatter-6-4.html" class="internal"><span class="codenumber">E.2</span> <span class="title">Numbering in the Back Matter</span></a></div></li>
</ul>
</li>
<li class="toc-item toc-appendix"><div class="toc-title-box"><a href="appendix-gfdl.html" class="internal"><span class="codenumber">F</span> <span class="title">GNU Free Documentation License</span></a></div></li>
<li class="toc-item toc-index"><div class="toc-title-box"><a href="backmatter-8.html" class="internal"><span class="title">Index</span></a></div></li>
<li class="toc-item toc-colophon"><div class="toc-title-box"><a href="back-colophon.html" class="internal"><span class="title">Colophon</span></a></div></li>
</ul>
</li>
</ul></nav></div>
<main class="ptx-main"><div id="ptx-content" class="ptx-content"><section class="section" id="activecode"><h2 class="heading hide-type">
<span class="type">Section</span><span class="space"> </span><span class="codenumber">5.2</span><span class="space"> </span><span class="title">ActiveCode</span>
</h2>
<div class="para" id="activecode-2">Programs in supported languages are made interactive in HTML when <code class="code-inline tex2jax_ignore">@interactive</code> is set to <code class="code-inline tex2jax_ignore">activecode</code>. Some languages can be made interactive on any server, while others require being served from Runestone servers. See <a class="external" href="https://pretextbook.org/doc/guide/html/topic-program-console.html#interactive-program-capabilities" target="_blank">Interactive Programs Capabilities in the PreTeXt Guide</a><details class="ptx-footnote" aria-live="polite" id="activecode-2-4"><summary class="ptx-footnote__number" title="Footnote 5.2.1"><sup> 1 </sup></summary><div class="ptx-footnote__contents" id="activecode-2-4"><code class="code-inline tex2jax_ignore">pretextbook.org/doc/guide/html/topic-program-console.html#interactive-program-capabilities</code></div></details> for a list of what languages are supported in which environs.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-2" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-python"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.1<span class="period">.</span></span><span class="space"> </span>An interactive Python program, using <span class="booktitle">Runestone</span><div class="autopermalink" data-description="Listing 5.2.1: An interactive Python program, using Runestone"><a href="#program-activecode-python" title="Copy heading and permalink for Listing 5.2.1: An interactive Python program, using Runestone">🔗</a></div></figcaption><div class="listing__contents"><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-hello-world"><textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-hello-world_editor" data-question_label="" data-codelens="true" data-timelimit="25000">print("Hello, World!")
</textarea></div></div></div></div></figure><div class="para" id="activecode-4">Otherwise they are rendered as text with syntax coloring. Either way, if a <code class="code-inline tex2jax_ignore">@language</code> is not specified, <code class="code-inline tex2jax_ignore">docinfo/programs/@language</code> will be checked to determine what language to assume the code is written in.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-4" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-python2"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.2<span class="period">.</span></span><span class="space"> </span>An interactive Python program, using <span class="booktitle">Runestone</span><div class="autopermalink" data-description="Listing 5.2.2: An interactive Python program, using Runestone"><a href="#program-activecode-python2" title="Copy heading and permalink for Listing 5.2.2: An interactive Python program, using Runestone">🔗</a></div></figcaption><div class="listing__contents"><div class="code-box"><pre class="program"><code id="rs-hello-world2" class="language-py">print("Hello, World!")
</code></pre></div></div></figure><div class="para" id="activecode-6">A C program will only be interactive if hosted on a Runestone server.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-6" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-c"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.3<span class="period">.</span></span><span class="space"> </span>An C program, interactive on a <span class="booktitle">Runestone</span> server<div class="autopermalink" data-description="Listing 5.2.3: An C program, interactive on a Runestone server"><a href="#program-activecode-c" title="Copy heading and permalink for Listing 5.2.3: An C program, interactive on a Runestone server">🔗</a></div></figcaption><div class="listing__contents"><pre class="program"><code id="rs-c-hello-world" class="language-c">#include &lt;stdio.h&gt;

int main(void)
{
    puts("Hello, world!");
}
</code></pre></div></figure><div class="para" id="activecode-8">A Java program will only be interactive if hosted on a Runestone server.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-8" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-java"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.4<span class="period">.</span></span><span class="space"> </span>Informal Java “Hello, World”A Java program, interactive on a <span class="booktitle">Runestone</span> server<div class="autopermalink" data-description='Listing 5.2.4: Informal Java "Hello, World"A Java program, interactive on a Runestone server'><a href="#program-activecode-java" title='Copy heading and permalink for Listing 5.2.4: Informal Java "Hello, World"A Java program, interactive on a Runestone server'>🔗</a></div></figcaption><div class="listing__contents"><pre class="program"><code id="rs-java-hello-world" class="language-java">public class HelloWorld {
    public static void main(String[] args) {
        System.out.println("Hi world!");
    }
}
</code></pre></div></figure><div class="para" id="activecode-10">Javascript programs are made interactive in HTML, on request.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-10" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-javascript"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.5<span class="period">.</span></span><span class="space"> </span>An interactive JavaScript program, using <span class="booktitle">Runestone</span><div class="autopermalink" data-description="Listing 5.2.5: An interactive JavaScript program, using Runestone"><a href="#program-activecode-javascript" title="Copy heading and permalink for Listing 5.2.5: An interactive JavaScript program, using Runestone">🔗</a></div></figcaption><div class="listing__contents"><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-javascript-hello-world"><textarea data-lang="javascript" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-javascript-hello-world_editor" data-question_label="" data-codelens="false" data-timelimit="25000">document.write('Hello, world!');
</textarea></div></div></div></div></figure><div class="para" id="activecode-12">Some languages, like Java or C++, are only interactive when run on a Runestone server where the code can be compiled and run. Those languages can specify <code class="code-inline tex2jax_ignore">@compiler-args</code> and <code class="code-inline tex2jax_ignore">@linker-args</code> or <code class="code-inline tex2jax_ignore">@interpreter-args</code> as appropriate to the language. Default values for those options can be set in <code class="code-inline tex2jax_ignore">&lt;docinfo/programs&gt;</code> - any defaults set there will be used for any program that lacks the corresponding attribute.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-12" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="activecode-13"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.6<span class="period">.</span></span><span class="space"> </span>A C++ program with compiler-args and stdin<div class="autopermalink" data-description="Listing 5.2.6: A C++ program with compiler-args and stdin"><a href="#activecode-13" title="Copy heading and permalink for Listing 5.2.6: A C++ program with compiler-args and stdin">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program interactive="activecode" label="hello-world-cpp" language="cpp" compiler-args="-std=c++17,-Wall,-g" linker-args="-g,-s" include-source="yes"&gt;
    &lt;code&gt;
        #include &amp;lt;iostream&amp;gt;
        #include &amp;lt;string&amp;gt;
        using namespace std;
        int main() {
            string name;
            cin &amp;gt;&amp;gt; name;
            cout &amp;lt;&amp;lt; "Hello, " &amp;lt;&amp;lt; name &amp;lt;&amp;lt; endl;
            return 0;
        }
    &lt;/code&gt;
    &lt;stdin&gt;Jane&lt;/stdin&gt;
&lt;/program&gt;
</code></pre></article></details><pre class="program"><code id="rs-hello-world-cpp" class="language-cpp">#include &lt;iostream&gt;
#include &lt;string&gt;
using namespace std;
int main() {
    string name;
    cin &gt;&gt; name;
    cout &lt;&lt; "Hello, " &lt;&lt; name &lt;&lt; endl;
    return 0;
}
</code></pre>
</div></figure><figure class="listing figure-like" id="program-static-java"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.7<span class="period">.</span></span><span class="space"> </span>A static Java program<div class="autopermalink" data-description="Listing 5.2.7: A static Java program"><a href="#program-static-java" title="Copy heading and permalink for Listing 5.2.7: A static Java program">🔗</a></div></figcaption><div class="listing__contents"><div class="code-box"><pre class="program"><code id="rs-java-static" class="language-java">import javax.swing.JFrame;  //Importing class JFrame
import javax.swing.JLabel;  //Importing class JLabel
public class HelloWorld {
    public static void main(String[] args) {
        JFrame frame = new JFrame();           //Creating frame
        frame.setTitle("Hi!");                 //Setting title frame
        frame.add(new JLabel("Hello, world!"));//Adding text to frame
        frame.pack();                          //Setting size to smallest
        frame.setLocationRelativeTo(null);     //Centering frame
        frame.setVisible(true);                //Showing frame
    }
}
</code></pre></div></div></figure><div class="para" id="activecode-15">An Octave program will also only be interactive if hosted on a Runestone server.  Octave is meant to be a drop-in replacement for Matlab.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-15" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-octave"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.8<span class="period">.</span></span><span class="space"> </span>A simple Octave program<div class="autopermalink" data-description="Listing 5.2.8: A simple Octave program"><a href="#program-activecode-octave" title="Copy heading and permalink for Listing 5.2.8: A simple Octave program">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program xml:id="octave-simple" interactive="activecode" language="octave" include-source="yes"&gt;
    &lt;code&gt;
    x = 2 + 2
    printf("%d\n", x)
    &lt;/code&gt;
&lt;/program&gt;
</code></pre></article></details><pre class="program"><code id="rs-octave-simple" class="language-matlab">x = 2 + 2
printf("%d\n", x)
</code></pre>
</div></figure><div class="para" id="activecode-17">A language not supported by Runestone Services will always be rendered static.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-17" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-mistake-pascal"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.9<span class="period">.</span></span><span class="space"> </span>A Pascal program that cannot be interactive on Runestone<div class="autopermalink" data-description="Listing 5.2.9: A Pascal program that cannot be interactive on Runestone"><a href="#program-mistake-pascal" title="Copy heading and permalink for Listing 5.2.9: A Pascal program that cannot be interactive on Runestone">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program xml:id="pascal-mistake" interactive="activecode" language="pascal" include-source="yes"&gt;
    &lt;code&gt;
    program HelloWorld;
    begin
      WriteLn('Hello, world!');
    end.
    &lt;/code&gt;
&lt;/program&gt;
</code></pre></article></details><pre class="program"><code id="rs-pascal-mistake" class="language-pascal">program HelloWorld;
begin
  WriteLn('Hello, world!');
end.
</code></pre>
</div></figure><div class="para" id="activecode-19">The <code class="code-inline tex2jax_ignore">@highlight-lines</code> works on ActiveCode programs, but the highlighted lines are only shown when viewing the initial code. Any version of the code that the reader has edited may have shifted lines of code and thus highlighting might affect the wrong lines, causing confusion.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-19" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program interactive="activecode" language="python" label="program-activecode-highlights" highlight-lines="1-2,5" include-source="yes"&gt;
    def add(a, b):
        return a + b

    # Use the function
    result = add(2, 3)
    if result == 5:
        print("Test passed")
    else:
        print("Test failed")
&lt;/program&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-program-activecode-highlights"><textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-program-activecode-highlights_editor" data-question_label="" data-codelens="true" data-timelimit="25000" data-highlight-lines="1-2,5">def add(a, b):
    return a + b

# Use the function
result = add(2, 3)
if result == 5:
    print("Test passed")
else:
    print("Test failed")
</textarea></div></div></div>
<div class="para" id="activecode-21">A program can have a <code class="code-inline tex2jax_ignore">&lt;preamble&gt;</code> and/or <code class="code-inline tex2jax_ignore">&lt;postamble&gt;</code> which are added to the code that the user writes before it is run. They are visible by default, but can be made invisible with <code class="code-inline tex2jax_ignore">@visible</code> set to <code class="code-inline tex2jax_ignore">"no"</code>. When visible, the code editor will prevent those regions from being modified. The indentation for lines in the <code class="code-inline tex2jax_ignore">preamble/code/postamble</code> elements will be calculated relative to each other - make sure to indent them all to a similar extent. (In the source for the sample below, the <code class="code-inline tex2jax_ignore"># TODO...</code> is intentially indented one stop extra so that the user’s code is part of the <code class="code-inline tex2jax_ignore">add</code> function.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-21" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-22">
<code class="code-inline tex2jax_ignore">&lt;tests&gt;</code> are similar to <code class="code-inline tex2jax_ignore">&lt;postamble&gt;</code> in that it represents code that is added to the users submission. However, <code class="code-inline tex2jax_ignore">&lt;tests&gt;</code> is intended specifically for unit testing code (see examples below for unit testing in Python, Java, C++, SQL). Tests are invisble by default and can be made visible with <code class="code-inline tex2jax_ignore">@visible</code> set to <code class="code-inline tex2jax_ignore">"yes"</code>. For historical reasons, the indentation of the <code class="code-inline tex2jax_ignore">&lt;tests&gt;</code> is treated separately from the rest of the program.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-22" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="activecode-23"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.10<span class="period">.</span></span><span class="space"> </span>A Python program with preamble/postamble<div class="autopermalink" data-description="Listing 5.2.10: A Python program with preamble/postamble"><a href="#activecode-23" title="Copy heading and permalink for Listing 5.2.10: A Python program with preamble/postamble">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program interactive="activecode" language="python" label="program-activecode-python-ambles" include-source="yes"&gt;
    &lt;preamble&gt;
    def add(a, b):
    &lt;/preamble&gt;
    &lt;code&gt;
        # TODO - complete the add function

    &lt;/code&gt;
    &lt;postamble&gt;
    # Use the function
    result = add(2, 3)
    if result == 5:
        print("Test passed")
    else:
        print("Test failed")
    &lt;/postamble&gt;
&lt;/program&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-program-activecode-python-ambles"><textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-program-activecode-python-ambles_editor" data-question_label="" data-codelens="true" data-timelimit="25000">def add(a, b):
^^^!
    # TODO - complete the add function

===!
# Use the function
result = add(2, 3)
if result == 5:
    print("Test passed")
else:
    print("Test failed")
</textarea></div></div></div>
</div></figure><div class="para" id="activecode-24">Here is the same Python program from the previous section, but now with a <code class="code-inline tex2jax_ignore">&lt;preamble&gt;</code> and <code class="code-inline tex2jax_ignore">&lt;postamble&gt;</code> that are invisible.  The user will not see the code that is added to their submission.  Not actually useful in this case, but it might be if you wanted to hide boilerplate setup from the reader.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-24" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="activecode-25"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.11<span class="period">.</span></span><span class="space"> </span>A Python program with invisible pre/post ambles<div class="autopermalink" data-description="Listing 5.2.11: A Python program with invisible pre/post ambles"><a href="#activecode-25" title="Copy heading and permalink for Listing 5.2.11: A Python program with invisible pre/post ambles">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program interactive="activecode" language="python" label="program-activecode-python-ambles-invisible" include-source="yes"&gt;
    &lt;preamble visible="no"&gt;
    def add(a, b):
    &lt;/preamble&gt;
    &lt;code&gt;
        # TODO - complete the add function
    &lt;/code&gt;
    &lt;postamble visible="no"&gt;

    # Use the function
    result = add(2, 3)
    if result == 5:
        print("Test passed")
    else:
        print("Test failed")
    &lt;/postamble&gt;
&lt;/program&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-program-activecode-python-ambles-invisible"><textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-program-activecode-python-ambles-invisible_editor" data-question_label="" data-codelens="true" data-timelimit="25000">def add(a, b):
^^^^
    # TODO - complete the add function
====

# Use the function
result = add(2, 3)
if result == 5:
    print("Test passed")
else:
    print("Test failed")
</textarea></div></div></div>
</div></figure><div class="para" id="activecode-26">The following Python program is in a <code class="code-inline tex2jax_ignore">&lt;listing&gt;</code> since we will want to reference it shortly.  The program does not do very much, it just defines four variables whose values are lists of statistics.  It should run, and there will be no syntax errors, but it is a bit boring since there is no output. Note that it does not have an <code class="code-inline tex2jax_ignore">@language</code> and is relying on the default one specified in <code class="code-inline tex2jax_ignore">&lt;docinfo/programs&gt;</code><div class="autopermalink" data-description="Paragraph"><a href="#activecode-26" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="listing-python-included"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.12<span class="period">.</span></span><span class="space"> </span>A Python program that defines some statistics<div class="autopermalink" data-description="Listing 5.2.12: A Python program that defines some statistics"><a href="#listing-python-included" title="Copy heading and permalink for Listing 5.2.12: A Python program that defines some statistics">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program xml:id="python-statistics" interactive="activecode" label="statistics" include-source="yes"&gt;
    &lt;code&gt;
    loan_amount = [1250.0, 500.0, 1450.0, 200.0, 700.0, 100.0, 250.0, 225.0, 1200.0, 150.0, 600.0, 300.0, 700.0, 125.0, 650.0, 175.0, 1800.0, 1525.0, 575.0, 700.0, 1450.0, 400.0, 200.0, 1000.0, 350.0]

    country_name = ['Azerbaijan', 'El Salvador', 'Bolivia', 'Paraguay', 'El Salvador', 'Philippines', 'Philippines', 'Nicaragua', 'Guatemala', 'Philippines', 'Paraguay', 'Philippines', 'Bolivia', 'Philippines', 'Philippines', 'Madagascar', 'Georgia', 'Uganda', 'Kenya', 'Tajikistan', 'Jordan', 'Kenya', 'Philippines', 'Ecuador', 'Kenya']

    time_to_raise = [193075.0, 1157108.0, 1552939.0, 244945.0, 238797.0, 1248909.0, 773599.0, 116181.0, 2288095.0, 51668.0, 26717.0, 48030.0, 1839190.0, 71117.0, 580401.0, 800427.0, 1156218.0, 1166045.0, 2924705.0, 470622.0, 24078.0, 260044.0, 445938.0, 201408.0, 2370450.0]

    num_lenders_total = [38, 18, 51, 3, 21, 1, 10, 8, 42, 1, 18, 6, 28, 5, 16, 7, 54, 1, 18, 22, 36, 12, 8, 24, 8]
    &lt;/code&gt;
&lt;/program&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-statistics"><textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-statistics_editor" data-question_label="" data-codelens="true" data-timelimit="25000">loan_amount = [1250.0, 500.0, 1450.0, 200.0, 700.0, 100.0, 250.0, 225.0, 1200.0, 150.0, 600.0, 300.0, 700.0, 125.0, 650.0, 175.0, 1800.0, 1525.0, 575.0, 700.0, 1450.0, 400.0, 200.0, 1000.0, 350.0]

country_name = ['Azerbaijan', 'El Salvador', 'Bolivia', 'Paraguay', 'El Salvador', 'Philippines', 'Philippines', 'Nicaragua', 'Guatemala', 'Philippines', 'Paraguay', 'Philippines', 'Bolivia', 'Philippines', 'Philippines', 'Madagascar', 'Georgia', 'Uganda', 'Kenya', 'Tajikistan', 'Jordan', 'Kenya', 'Philippines', 'Ecuador', 'Kenya']

time_to_raise = [193075.0, 1157108.0, 1552939.0, 244945.0, 238797.0, 1248909.0, 773599.0, 116181.0, 2288095.0, 51668.0, 26717.0, 48030.0, 1839190.0, 71117.0, 580401.0, 800427.0, 1156218.0, 1166045.0, 2924705.0, 470622.0, 24078.0, 260044.0, 445938.0, 201408.0, 2370450.0]

num_lenders_total = [38, 18, 51, 3, 21, 1, 10, 8, 42, 1, 18, 6, 28, 5, 16, 7, 54, 1, 18, 22, 36, 12, 8, 24, 8]
</textarea></div></div></div>
</div></figure><div class="para" id="activecode-28">An ActiveCode running on the Runestone server (rather than in a browser) can be provided with attributes that are flags to influence the compiler and linker, as in this <code class="code-inline tex2jax_ignore">&lt;program&gt;</code> recycled from before.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-28" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="program-activecode-java-flags"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.13<span class="period">.</span></span><span class="space"> </span>A Java program, interactive on a <span class="booktitle">Runestone</span> server, with compiler and linker flags<div class="autopermalink" data-description="Listing 5.2.13: A Java program, interactive on a Runestone server, with compiler and linker flags"><a href="#program-activecode-java-flags" title="Copy heading and permalink for Listing 5.2.13: A Java program, interactive on a Runestone server, with compiler and linker flags">🔗</a></div></figcaption><div class="listing__contents"><pre class="program"><code id="rs-java-hello-world-flags" class="language-java">import javax.swing.JFrame;  //Importing class JFrame
import javax.swing.JLabel;  //Importing class JLabel
public class HelloWorld {
    public static void main(String[] args) {
        JFrame frame = new JFrame();           //Creating frame
        frame.setTitle("Hi!");                 //Setting title frame
        frame.add(new JLabel("Hello, world!"));//Adding text to frame
        frame.pack();                          //Setting size to smallest
        frame.setLocationRelativeTo(null);     //Centering frame
        frame.setVisible(true);                //Showing frame
    }
}
</code></pre></div></figure><div class="para" id="activecode-30">Now a programming exercise.  The program upcoming is going to <dfn class="terminology">include</dfn> all the code of the program preceding.  This is accomplished with an <code class="code-inline tex2jax_ignore">@include</code> attribute on the <em class="emphasis">including</em> program whose value is the <code class="code-inline tex2jax_ignore">@xml:id</code> of the <em class="emphasis">included</em> program.  So by running the next program, it should pass all of its three tests (for example another example using unit tests, see <a href="coding-exercises.html#coding-exercise-python-unit-test" class="xref" data-knowl="./knowl/xref/unit-test-example.html" data-reveal-label="Reveal" data-close-label="Close" title="Checkpoint 5.4.4: Coding Exercise, with Unit Tests">Checkpoint 5.4.4</a>).  Now reload the page, do not run the program in the listing, and then see that the program in the exercise still runs correctly.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-30" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-31">You’ll see nothing that tells the reader that the one chunk of code is prefacing the other.  And in static formats it might be even less obvious.  So you will want to say <em class="emphasis">something</em> to alert the reader.  Here it is easy: <a href="activecode.html#exercise-python-including" class="xref" data-knowl="./knowl/xref/exercise-python-including.html" data-reveal-label="Reveal" data-close-label="Close" title="Checkpoint 5.2.14: A Python program, including another">Checkpoint 5.2.14</a> includes all the code from <a href="activecode.html#listing-python-included" class="xref" data-knowl="./knowl/xref/listing-python-included.html" data-reveal-label="Reveal" data-close-label="Close" title="Listing 5.2.12: A Python program that defines some statistics">Listing 5.2.12</a>.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-31" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-32">This program also makes use of <code class="code-inline tex2jax_ignore">@autorun</code> to execute on page load and the <code class="code-inline tex2jax_ignore">@codelens</code>to disable the codelens feature.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-32" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<article class="exercise exercise-like" id="exercise-python-including"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.2.14</span><span class="period">.</span><span class="space"> </span><span class="title">A Python program, including another.</span>
</h3>
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for exercise</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;exercise xml:id="exercise-python-including" include-source="yes"&gt;
    &lt;title&gt;A Python program, including another&lt;/title&gt;
    &lt;statement&gt;
        &lt;p&gt;Compute the total amount of money loaned and store it in the variable &lt;c&gt;loan_total&lt;/c&gt;.&lt;/p&gt;
    &lt;/statement&gt;
    &lt;program autorun="yes" codelens="no" xml:id="python-summation" interactive="activecode" language="python" label="python-sum-total" include="python-statistics"&gt;
        &lt;code&gt;
        loan_total = 0
        for loan in loan_amount:
            loan_total += loan
        print(loan_total)
        &lt;/code&gt;
        &lt;tests&gt;
        from unittest.gui import TestCaseGui

        class MyTests(TestCaseGui):

            def testOne(self):
                self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
                self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
                self.assertFalse('sum(' in self.getEditorText(), "you may not use sum()")

        MyTests().main()
        &lt;/tests&gt;
    &lt;/program&gt;
&lt;/exercise&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-exercise-python-including">
<div class="ac_question exercise-statement" id="rs-exercise-python-including_question"><div class="para" id="exercise-python-including-2-1">Compute the total amount of money loaned and store it in the variable <code class="code-inline tex2jax_ignore">loan_total</code>.<div class="autopermalink" data-description="Paragraph"><a href="#exercise-python-including-2-1" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div></div>
<textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-exercise-python-including_editor" data-question_label="" data-codelens="false" data-include="rs-statistics" data-timelimit="25000" data-autorun="yes">loan_total = 0
for loan in loan_amount:
    loan_total += loan
print(loan_total)
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
        self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
        self.assertFalse('sum(' in self.getEditorText(), "you may not use sum()")

MyTests().main()
</textarea>
</div></div></div>
<div class="autopermalink" data-description="Checkpoint 5.2.14: A Python program, including another"><a href="#exercise-python-including" title="Copy heading and permalink for Checkpoint 5.2.14: A Python program, including another">🔗</a></div></article><div class="para" id="activecode-34">Exact same exercise again, but now we include <em class="emphasis">two</em> programs.  We first get the simple “Hello, world!” program at <a href="activecode.html#program-activecode-python" class="xref" data-knowl="./knowl/xref/program-activecode-python.html" data-reveal-label="Reveal" data-close-label="Close" title="Listing 5.2.1: An interactive Python program, using Runestone">Listing 5.2.1</a> and then the same program defining the variables with lists of statistics at <a href="activecode.html#listing-python-included" class="xref" data-knowl="./knowl/xref/listing-python-included.html" data-reveal-label="Reveal" data-close-label="Close" title="Listing 5.2.12: A Python program that defines some statistics">Listing 5.2.12</a>.  So the output just includes the extra result from the <code class="code-inline tex2jax_ignore">print()</code> statement.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-34" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-35">This program also makes use of <code class="code-inline tex2jax_ignore">@hidecode</code> to initially keep the code hidden and <code class="code-inline tex2jax_ignore">@download</code> to enable a file download of the program (that includes all the included code).<div class="autopermalink" data-description="Paragraph"><a href="#activecode-35" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<article class="exercise exercise-like" id="exercise-python-including-two"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.2.15</span><span class="period">.</span><span class="space"> </span><span class="title">A Python program, including two others.</span>
</h3>
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for exercise</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;exercise xml:id="exercise-python-including-two" include-source="yes"&gt;
    &lt;title&gt;A Python program, including two others&lt;/title&gt;
    &lt;statement&gt;
        &lt;p&gt;Compute the total amount of money loaned and store it in the variable &lt;c&gt;loan_total&lt;/c&gt;.&lt;/p&gt;
    &lt;/statement&gt;
    &lt;program xml:id="python-summation-two" interactive="activecode" language="python" label="python-sum-total-two" include="python-hello-world python-statistics " hidecode="yes" download="yes"&gt;
        &lt;code&gt;
        loan_total = 0
        for loan in loan_amount:
            loan_total += loan
        print(loan_total)
        &lt;/code&gt;
        &lt;tests&gt;
        from unittest.gui import TestCaseGui

        class MyTests(TestCaseGui):

            def testOne(self):
                self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
                self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
                self.assertFalse('sum(' in self.getEditorText(), "you may not use sum")

        MyTests().main()
        &lt;/tests&gt;
    &lt;/program&gt;
&lt;/exercise&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-exercise-python-including-two">
<div class="ac_question exercise-statement" id="rs-exercise-python-including-two_question"><div class="para" id="exercise-python-including-two-2-1">Compute the total amount of money loaned and store it in the variable <code class="code-inline tex2jax_ignore">loan_total</code>.<div class="autopermalink" data-description="Paragraph"><a href="#exercise-python-including-two-2-1" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div></div>
<textarea data-lang="python" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-exercise-python-including-two_editor" data-question_label="" data-codelens="true" data-include="rs-hello-world rs-statistics" data-timelimit="25000" data-enabledownload="yes" data-hidecode="yes">loan_total = 0
for loan in loan_amount:
    loan_total += loan
print(loan_total)
====
from unittest.gui import TestCaseGui

class MyTests(TestCaseGui):

    def testOne(self):
        self.assertTrue('loan_total' in self.getEditorText(), "you need a loan_total variable")
        self.assertEqual(loan_total, sum(loan_amount), "Use the accumulator pattern to add up all the loans")
        self.assertFalse('sum(' in self.getEditorText(), "you may not use sum")

MyTests().main()
</textarea>
</div></div></div>
<div class="autopermalink" data-description="Checkpoint 5.2.15: A Python program, including two others"><a href="#exercise-python-including-two" title="Copy heading and permalink for Checkpoint 5.2.15: A Python program, including two others">🔗</a></div></article><div class="para" id="activecode-37">Here is an activecode with <code class="code-inline tex2jax_ignore">@language</code> set to <code class="code-inline tex2jax_ignore">sql</code> uses the <code class="code-inline tex2jax_ignore">@database</code> to load a SQLite database file.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-37" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<article class="exercise exercise-like" id="exercise-sql-using-db"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.2.16</span><span class="period">.</span><span class="space"> </span><span class="title">An SQL program that uses an SQLite database file.</span>
</h3>
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for exercise</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;exercise xml:id="exercise-sql-using-db" include-source="yes"&gt;
    &lt;title&gt;An SQL program that uses an SQLite database file&lt;/title&gt;
    &lt;statement&gt;
        &lt;p&gt;Select all the columns of all the rows in the &lt;c&gt;test&lt;/c&gt; database table.&lt;/p&gt;
    &lt;/statement&gt;
    &lt;program xml:id="sql-using-db" interactive="activecode" language="sql" label="sql-using-db" database="sql/test.db"&gt;
        &lt;code&gt;
            SELECT * FROM test
        &lt;/code&gt;
        &lt;tests&gt;
            assert 1,1 == world
            assert 0,1 == hello
            assert 2,1 == 35
        &lt;/tests&gt;
    &lt;/program&gt;
&lt;/exercise&gt;
</code></pre></article></details><div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="rs-exercise-sql-using-db">
<div class="ac_question exercise-statement" id="rs-exercise-sql-using-db_question"><div class="para" id="exercise-sql-using-db-2-1">Select all the columns of all the rows in the <code class="code-inline tex2jax_ignore">test</code> database table.<div class="autopermalink" data-description="Paragraph"><a href="#exercise-sql-using-db-2-1" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div></div>
<textarea data-lang="sql" data-audio="" data-coach="true" style="visibility: hidden;" id="rs-exercise-sql-using-db_editor" data-question_label="" data-codelens="false" data-wasm="/_static" data-dburl="external/sql/test.db" data-timelimit="25000">SELECT * FROM test
====
assert 1,1 == world
assert 0,1 == hello
assert 2,1 == 35
</textarea>
</div></div></div>
<div class="autopermalink" data-description="Checkpoint 5.2.16: An SQL program that uses an SQLite database file"><a href="#exercise-sql-using-db" title="Copy heading and permalink for Checkpoint 5.2.16: An SQL program that uses an SQLite database file">🔗</a></div></article><div class="para" id="activecode-39">Now a C++ program that is spread across multiple files. This example will only be interactive on a Runestone server.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-39" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-40">First, we have a <code class="code-inline tex2jax_ignore">.h</code> file that defines an <code class="code-inline tex2jax_ignore">@filename</code>. It has an <code class="code-inline tex2jax_ignore">@xml:id</code> so other elements can reference it. It also has an <code class="code-inline tex2jax_ignore">@label</code> which is it’s unique identifier in the Runestone database. (The label is intentionally different than the xml:id for demonstration purposes, but they can be the same.) The <code class="code-inline tex2jax_ignore">@filename</code> is used to indicate what to name to use when the contents of the element are written to a file on the server - other program code and/or tools on the server can look for it using this name. The <code class="code-inline tex2jax_ignore">@filename</code> need not be unique—we could have multiple programs with <code class="code-inline tex2jax_ignore">@filename="add.h"</code>. As we will see below, the author specifies which “version” of a file should be used in any given location with its unique xml:id.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-40" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-41">The <code class="code-inline tex2jax_ignore">&lt;program&gt;</code> is inside of a <code class="code-inline tex2jax_ignore">&lt;listing&gt;</code> so that we have a place to add a caption and so that we can reference the code sample, including that caption, with an <code class="code-inline tex2jax_ignore">&lt;xref&gt;</code> from elsewhere. Note that the <code class="code-inline tex2jax_ignore">&lt;listing&gt;</code> has both a <code class="code-inline tex2jax_ignore">&lt;caption&gt;</code> and a <code class="code-inline tex2jax_ignore">&lt;title&gt;</code>. The <code class="code-inline tex2jax_ignore">&lt;caption&gt;</code> will be rendered with the program. The <code class="code-inline tex2jax_ignore">&lt;title&gt;</code> will not appear locally. It will be used if an <code class="code-inline tex2jax_ignore">&lt;xref&gt;</code> links to the <code class="code-inline tex2jax_ignore">&lt;listing&gt;</code> using <code class="code-inline tex2jax_ignore">@text="title"</code><div class="autopermalink" data-description="Paragraph"><a href="#activecode-41" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<figure class="listing figure-like" id="addh-v1-listing"><figcaption><span class="type">Listing</span><span class="space"> </span><span class="codenumber">5.2.17<span class="period">.</span></span><span class="space"> </span>add.h (version 1)<div class="autopermalink" data-description="Listing 5.2.17: add.h (version 1)"><a href="#addh-v1-listing" title="Copy heading and permalink for Listing 5.2.17: add.h (version 1)">🔗</a></div></figcaption><div class="listing__contents">
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program filename="add.h" xml:id="addh-v1" label="addh-v1-label" language="cpp" include-source="yes"&gt;
    int add(int a, int b);
&lt;/program&gt;
</code></pre></article></details><div class="code-box"><pre class="program"><code id="rs-addh-v1-label" data-filename="add.h" class="language-cpp">int add(int a, int b);
</code></pre></div>
</div></figure><div class="para" id="activecode-43">Next, a <code class="code-inline tex2jax_ignore">.cpp</code> file. Note that because it is part of an exercise, the <code class="code-inline tex2jax_ignore">@label</code> is applied to the exercise that contains it. The <code class="code-inline tex2jax_ignore">@xml:id</code> and <code class="code-inline tex2jax_ignore">@filename</code> still belong on the program itself.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-43" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-44">It also has <code class="code-inline tex2jax_ignore">@extra-compiler-args</code>. These will be added to any default compiler args for the book. In this case, the <code class="code-inline tex2jax_ignore">-c</code> indicates that we only want to compile this file when it is “Run” and not try to link it or actually run it.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-44" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<article class="exercise exercise-like" id="addcpp-v1"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.2.18</span><span class="period">.</span>
</h3>
<div class="para" id="addcpp-v1-1-1">You can leave this code as is or modify it. When you click “Run”, the code will be compiled. However, it will not be run as this is not a standalone program. To run it, use the full program below.<div class="autopermalink" data-description="Paragraph"><a href="#addcpp-v1-1-1" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<pre class="program"><code id="rs-addcpp-v1" data-filename="add.cpp" class="language-cpp">int add(int a, int b) {
    return a + b;
}
</code></pre>
<div class="autopermalink" data-description="Checkpoint 5.2.18"><a href="#addcpp-v1" title="Copy heading and permalink for Checkpoint 5.2.18">🔗</a></div></article><div class="para" id="activecode-46">Finally, the core program. <code class="code-inline tex2jax_ignore">@add-files</code> is used to specify (by xml:id) the files that need to be added to the program directory. <code class="code-inline tex2jax_ignore">@compile-also</code> specifies files that must be compiled with this source file (they will be assumed to be also part of <code class="code-inline tex2jax_ignore">@add-files</code>). Notice that when using <code class="code-inline tex2jax_ignore">@add-files</code> to reference code we want to include, we use the xml:id of the target program - it has the contents we want to include. When making a textual reference with an <code class="code-inline tex2jax_ignore">&lt;xref&gt;</code> we can’t link to the program, as it lacks the contextual information required to create a valid reference. Instead, we should link to a listing (or some other container) that surrounds the code, like this: <a href="activecode.html#addh-v1-listing" class="xref" data-knowl="./knowl/xref/addh-v1-listing.html" data-reveal-label="Reveal" data-close-label="Close" title="Listing 5.2.17: add.h (version 1)">Listing 5.2.17</a> or <a href="activecode.html#addh-v1-listing" class="xref" data-knowl="./knowl/xref/addh-v1-listing.html" data-reveal-label="Reveal" data-close-label="Close" title="Listing 5.2.17: add.h (version 1)">add.h (version 1)</a><div class="autopermalink" data-description="Paragraph"><a href="#activecode-46" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<details class="source-view"><summary class="source-view__link" data-reveal-label="Open" data-close-label="Close">View Source for program</summary><article class="view-source-view__content"><pre><code class="language-xml">&lt;program label="multiple-file-program" interactive="activecode" add-files="addh-v1" compile-also="addcpp-v1" include-source="yes" language="cpp"&gt;
    
    #include "add.h"
    #include &amp;lt;iostream&amp;gt;
    using namespace std;
    int main() {
        int a = 1;
        int b = 2;
        cout &amp;lt;&amp;lt; "The sum of " &amp;lt;&amp;lt; a &amp;lt;&amp;lt; " and " &amp;lt;&amp;lt; b &amp;lt;&amp;lt; " is " &amp;lt;&amp;lt; add(a, b) &amp;lt;&amp;lt; endl;
    }
&lt;/program&gt;
</code></pre></article></details><pre class="program"><code id="rs-multiple-file-program" class="language-cpp">#include "add.h"
#include &lt;iostream&gt;
using namespace std;
int main() {
    int a = 1;
    int b = 2;
    cout &lt;&lt; "The sum of " &lt;&lt; a &lt;&lt; " and " &lt;&lt; b &lt;&lt; " is " &lt;&lt; add(a, b) &lt;&lt; endl;
}
</code></pre>
<div class="para" id="activecode-48">Note that there is a cross page test of <code class="code-inline tex2jax_ignore">@add-files</code> located in <a href="runestone-hodgepodge.html" class="internal" title="Exercises 5.20: Hodgepodge">Exercises 5.20</a><div class="autopermalink" data-description="Paragraph"><a href="#activecode-48" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="para" id="activecode-49">A nonsense paragraph just to check on spacing.  A nonsense paragraph just to check on spacing.  A nonsense paragraph just to check on spacing.  A nonsense paragraph just to check on spacing.<div class="autopermalink" data-description="Paragraph"><a href="#activecode-49" title="Copy heading and permalink for Paragraph">🔗</a></div>
</div>
<div class="autopermalink" data-description="Section 5.2: ActiveCode"><a href="#activecode" title="Copy heading and permalink for Section 5.2: ActiveCode">🔗</a></div></section></div>
<div id="ptx-content-footer" class="ptx-content-footer">
<a class="previous-button button" href="inactive-code.html" title="Previous"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5cb;</span><span class="name">Prev</span></a><a class="top-button button" href="#" title="Top"><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5ce;</span><span class="name">Top</span></a><a class="next-button button" href="code-lens.html" title="Next"><span class="name">Next</span><span class="icon material-symbols-outlined" aria-hidden="true">&#xe5cc;</span></a>
</div></main>
</div>
<div id="ptx-page-footer" class="ptx-page-footer">
<a class="pretext-link" href="https://pretextbook.org" title="PreTeXt"><div class="logo"><svg xmlns="http://www.w3.org/2000/svg" height="100%" viewBox="338 3000 8772 6866" role="img"><title>PreTeXt logo</title><g style="stroke-width:.025in; stroke:currentColor; fill:none"><polyline points="472,3590 472,9732 " style="stroke-width:174; stroke-linejoin:miter; stroke-linecap:round; "></polyline><path style="stroke-width:126;stroke-linecap:butt;" d="M 4724,9448 A 4660 4660  0  0  1  8598  9259"></path><path style="stroke-width:174;stroke-linecap:butt;" d="M 4488,9685 A 4228 4228  0  0  0   472  9732"></path><path style="stroke-width:126;stroke-linecap:butt;" d="M 4724,3590 A 4241 4241  0  0  1  8598  3496"></path><path style="stroke-width:126;stroke-linecap:round;" d="M 850,3496  A 4241 4241  0  0  1  4724  3590"></path><path style="stroke-width:126;stroke-linecap:round;" d="M 850,9259  A 4507 4507  0  0  1  4724  9448"></path><polyline points="5385,4299 4062,8125" style="stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8598,3496 8598,9259" style="stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="850,3496 850,9259" style="stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="4960,9685 4488,9685" style="stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="3070,4582 1889,6141 3070,7700" style="stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="6418,4582 7600,6141 6418,7700" style="stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8976,3590 8976,9732" style="stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><path style="stroke-width:174;stroke-linecap:butt;" d="M 4960,9685 A 4228 4228  0  0  1  8976  9732"></path></g></svg></div></a><a class="runestone-link" href="https://runestone.academy" title="Runestone Academy"><img class="logo" src="https://runestone.academy/runestone/static/images/RAIcon_cropped.png" alt="Runstone Academy logo"></a><a class="mathjax-link" href="https://www.mathjax.org" title="MathJax"><img class="logo" src="https://www.mathjax.org/badge/badge-square-2.png" alt="MathJax logo"></a>
</div>
<!--Start: StatCounter code--><!--*** DO NOT COPY ANOTHER PROJECT'S MAGIC NUMBERS/ID ***--><!--***        GET YOUR OWN FROM STATCOUNTER           ***--><script>
var sc_project=9664088;
var sc_invisible=1;
var sc_security="6e60c510";
var scJsHost = (("https:" == document.location.protocol) ? "https://secure." : "https://www.");
document.write("<sc"+"ript src='" + scJsHost+ "statcounter.com/counter/counter.js'></"+"script>");
</script><noscript><div class="statcounter"><a title="web analytics" href="https://statcounter.com/" target="_blank"><img class="statcounter" src="https://c.statcounter.com/9664088/0/6e60c510/1/" alt="web analytics"></a></div></noscript>
<!--End: StatCounter code-->
<!--Start: Google Global Site Tag code--><!--*** DO NOT COPY ANOTHER PROJECT'S MAGIC NUMBERS/ID ***--><!--***           GET YOUR OWN FROM GOOGLE             ***-->
<script async="" src="https://www.googletagmanager.com/gtag/js?id=UA-48811660-1"></script><script>  window.dataLayer = window.dataLayer || [];
  function gtag(){dataLayer.push(arguments);}
  gtag('js', new Date());
  gtag('config', 'UA-48811660-1');
</script><!--End: Google Global Site Tag code-->
</body>
</html>
